Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Mohammad Noor AbuAsbe"
COLLABORATORS = ""
```

---

# Assignment 2: Applications of modular arithmetics

```python
import math
```

## Question 1: Cryptography

You will implement and analyze two simple cryptographic algorithms — Caesar's and Scytale ciphers.

1. Implement the encoder for Ceasar's cipher: https://en.wikipedia.org/wiki/Caesar_cipher (encode only the letters, leave other characters as-is).

Q1.1: Implement function `ceasar_cipher`, which will transform regular text into encoded text
<pre> 
------------
Input parameters:
- plaintext - string, the text that should be encrypted
- N - int, size of shift
------------
return value:
string, the text after encoding
</pre> 

<!-- grade_id: cell-fcc26ea16c7fcc6e, points: 0, grade: false, solution: true -->
```python
def ceasar_cipher(plaintext,N=1):    
   # YOUR CODE HERE
   ciphText = ""
   for char in plaintext:
     if char.isalpha():
       if ord('a') <= ord(char):
         ciphText = alphaReducer('a', char, N, ciphText)
       else:
         ciphText = alphaReducer('A', char, N, ciphText)
     else:
       ciphText = ciphText + char
   return ciphText

def alphaReducer(firstCharInSize, char, N, ciphText):
   val = ((ord(char) - ord(firstCharInSize)) + N) % 26
   val = val + ord(firstCharInSize)
   ciphText = ciphText + chr(val)
   return ciphText
```

<!-- grade_id: cell-910c1caba85f9efe, points: 10, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q1.1 --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'ceasarCipher' function..\n ")
assert type(ceasar_cipher("meep meep",1))==str,'output type is incorrect, should be string'
print ("good job!\nTests passed. There may be hidden tests too")

```

Output:
```
Testing the implementation of the 'ceasarCipher' function..
 
good job!
Tests passed. There may be hidden tests too
```

2. Implement a decoder function that can reverse the cypher you've created, and demonstrate it, by encoding and then decoding the implementation (the source code) of the Caesar cipher using Caesar cipher:
  * Encode the Python code 
  * Decode the result of encoding
  * Compare the original and the recovered code  - you can add a box where you print the results — they should be the same.

Q1.2A: Implement function `ceasar_cipher_decode`, which will restore the cyphertext created by `ceasar_cipher` to plaintext
<pre> 
------------
Input parameters:
- cyphertext - string, the encoded text that needs to be decyphered
- N - int, size of shift that was used in the encoding
------------
return value:
string, the text after decoding
</pre> 

<!-- grade_id: cell-fd561d08a5195124, points: 0, grade: false, solution: true -->
```python
def ceasar_cipher_decode(cypherText,N=1):
    # YOUR CODE HERE
    return ceasar_cipher(cypherText , -N)
```

```python
def TestForCeasar(originaltext,N):
    cyphered = ceasar_cipher(originaltext,N)
    deCyphered = ceasar_cipher_decode(cyphered,N)
    if deCyphered == originaltext:
      print("The text returned to his original State")
      print("The original is = " + originaltext)
      print("After Cyphiring it became = " + cyphered)
      print("After DeCyphiring it became = " + deCyphered)
    else:
      print("The text didnt return to its original State")

originaltext = "TH3 QUICK BRoWN F0X JUMPS OVER THE LAZY DOG, Yay!"
TestForCeasar(originaltext,23)
```

Output:
```
The text returned to his original State
The original is = TH3 QUICK BRoWN F0X JUMPS OVER THE LAZY DOG, Yay!
After Cyphiring it became = QE3 NRFZH YOlTK C0U GRJMP LSBO QEB IXWV ALD, Vxv!
After DeCyphiring it became = TH3 QUICK BRoWN F0X JUMPS OVER THE LAZY DOG, Yay!
```

<!-- grade_id: cell-a3e55a24e9359804, points: 8, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q1.2A --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'ceasarCipherDecode' function..\n ")
assert type(ceasar_cipher_decode("ABC",3))==str,'output type is incorrect, should be string'
print ("good job!\nTests passed. There may be hidden tests too")
```

Output:
```
Testing the implementation of the 'ceasarCipherDecode' function..
 
good job!
Tests passed. There may be hidden tests too
```

Q1.2B: Implement function `return_encrypted_ceasar_code`, which will return a string that equals the code of `ceasar_cipher` after encryption, with N=2 . Then, implement a function `return_decoded_ceasar_code` which will call `return_encrypted_ceasar_code` and return the decrypted outcome. 
<pre> 
return_encrypted_ceasar_code()
------------
Input parameters:
none
------------
return value:
string, the encrypted source code
</pre> 

<pre> 
return_decoded_ceasar_code()
------------
Input parameters:
encrypted_code - the output of return_encrypted_ceasar_code()
------------
return value:
string, the source code restored to plaintext
</pre> 

<!-- grade_id: cell-0c2fa8c5fcfb1b21, points: 0, grade: false, solution: true -->
```python
def return_encrypted_ceasar_code():
    # YOUR CODE HERE
    plaintext =  """ciphText = ""
   for char in plaintext:
     if char.isalpha():
       if ord('a') <= ord(char):
         ciphText = alphaReducer('a', char, N, ciphText)
       else:
         ciphText = alphaReducer('A', char, N, ciphText)
     else:
       ciphText = ciphText + char
   return ciphText

def alphaReducer(firstCharInSize, char, N, ciphText):
   val = ((ord(char) - ord(firstCharInSize)) + N) % 26
   val = val + ord(firstCharInSize)
   ciphText = ciphText + chr(val)
   return ciphText"""
    return ceasar_cipher(plaintext, 2)
    
def return_decoded_ceasar_code(encrypted_code=return_encrypted_ceasar_code()):
    # YOUR CODE HERE
    return ceasar_cipher_decode (encrypted_code, 2)
```

<!-- grade_id: cell-290384e75f4cb72f, points: 2, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q1.2B --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'ceasarCipherDecode' function..\n ")
assert type(return_encrypted_ceasar_code())==str,'output type is incorrect, should be string'
assert type(return_decoded_ceasar_code())==str,'output type is incorrect, should be string'
assert type(return_decoded_ceasar_code("some_arbitrary_string"))==str,'output type is incorrect, should be string'
print ("good job!\nTests passed. There may be hidden tests too")
```

Output:
```
Testing the implementation of the 'ceasarCipherDecode' function..
 
good job!
Tests passed. There may be hidden tests too
```


3. Find a real English word consisting of at least 3 letters which, when encoded using Caesar's cipher, is transformed to another real English word. 

Q1.3: Implement function `return_word_and_shift`, which returns a word in English, 3 letters or more in length, and an integer N so that if we apply Ceasar cypher to the word with shift N, we get another word in English.
<pre> 
------------
Input parameters:
none
------------
return value:
(word,shift)
- word - string, a word in English
- shift - integer, the size of N needed to encode it so that the cyphertext is also a word in English
</pre> 

<!-- grade_id: cell-61c68e5bc86ea006, points: 0, grade: false, solution: true -->
```python
def return_word_and_shift():
    # YOUR CODE HERE
    return ("Sleep", 9)
```

<!-- grade_id: cell-fcd88242abcbb20b, points: 10, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q1.3 --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'hash_int' function..\n ")
found_word=return_word_and_shift()
assert len(found_word)==2, 'function should return tuple with two values'
assert type(found_word[0]) is str, 'the first part of the output should be of type string'
assert type(found_word[1]) is int, 'the secondpart of the output should be of type int'
assert len(found_word[0])>=3, 'the word should be at least 3 letters long'
print ("good job!\nTests passed. There may be hidden tests too")


```

Output:
```
Testing the implementation of the 'hash_int' function..
 
good job!
Tests passed. There may be hidden tests too
```

4. Implement the encoder for the Scytale cipher: https://en.wikipedia.org/wiki/Scytale 
(encode all the characters, including spaces and non-letters).

Q1.4: Implement function `scytale_encode`, which will encrypt the input text according to the scytale code and with N as "cylinder width".
<pre> 
------------
Input parameters:
- plaintext - string, the text that needs to be encoded
- N - int, size of jump (for instance, in the wikipedia example, N=5). 
------------
return value:
string, the text after decoding
</pre> 

<!-- grade_id: cell-a74f19da20d8c32c, points: 0, grade: false, solution: true -->
```python
def scytale_encode(plaintext,N=5):
    # YOUR CODE HERE
    scytale = []
    encoded = ""
    i = 0
    k = 0
    while i < len(plaintext):
      loop = []
      while k < N and i < len(plaintext):
        loop.append(plaintext[i])
        i = i + 1
        k = k + 1
      scytale.append(loop)
      k = 0
    for k in range(N):
      for i in range(len(scytale)):
        if k < len(scytale[i]):
          encoded = encoded + scytale[i][k]
    return encoded
```

<!-- grade_id: cell-fede5aa2fb49e35f, points: 10, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q1.4 --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'ceasarCipherDecode' function..\n ")
assert type(scytale_encode("ABCabcABC",3))==str,'output type is incorrect, should be string'
print ("good job!\nTests passed. There may be hidden tests too")

```

Output:
```
Testing the implementation of the 'ceasarCipherDecode' function..
 
good job!
Tests passed. There may be hidden tests too
```


5. Implement a decoder, and demonstrate by encoding and then decoding the implementation of the Scytale filter, as before.

<b>Please note:</b> in the case where the encoded text length isn't a multiple of the jump N, the decoder should still be able to reverse the encoding correctly! Make sure you've handled this case.

Q1.5A: Implement function `scytale_cipher_decode`, which will restore the cyphertext created by `scytale_encode` to plaintext
<pre> 
------------
Input parameters:
- cyphertext - string, the encoded text that needs to be decyphered
- N - int, size of jump
------------
return value:
string, the text after decoding
</pre> 

<!-- grade_id: cell-1780c42a6fe02cf5, points: 0, grade: false, solution: true -->
```python
def scytale_cipher_decode(cyphertext,N):
    # YOUR CODE HERE
    scytale = []
    encoded = ""
    i = 0
    k = 0
    neededSize = len(cyphertext)/N
    extraSize = len(cyphertext) % N
    lastSize = 0
    while i < len(cyphertext):
      loop = []
      lastSize = neededSize
      if extraSize > 0:
        lastSize = lastSize + 1
        extraSize = extraSize - 1
      k = 0
      while k < int(lastSize) and i < len(cyphertext):
        loop.append(cyphertext[i])
        k = k + 1
        i = i + 1
      scytale.append(loop)
    for j in range(len(scytale[0])):
      for arr in range(len(scytale)):
        if j < len(scytale[arr]):
          encoded = encoded + scytale[arr][j]
    return encoded
```

<!-- grade_id: cell-7525afb85c29fb34, points: 8, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q1.5A --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'ceasarCipherDecode' function..\n ")
assert type(scytale_cipher_decode("ABCDefgh",4))==str,'output type is incorrect, should be string'
print ("good job!\nTests passed. There may be hidden tests too")
```

Output:
```
Testing the implementation of the 'ceasarCipherDecode' function..
 
good job!
Tests passed. There may be hidden tests too
```

Q1.5B: Implement functions `return_encrypted_scytale_code`, which will return a string that equals the code of `scytale_encode` after encryption, with N=3 . Then, implement a function `return_decoded_scytale_code` which will call `return_encrypted_scytale_code` and return the decrypted outcome. You may print the output to verify that it matches the original.
<pre> 
return_encrypted_scytale_code()
------------
Input parameters:
none
------------
return value:
string, the encrypted source code
</pre> 

<pre> 
return_decoded_scytale_code()
------------
Input parameters:
encrypted_code - the output of return_encrypted_scytale_code()
------------
return value:
string, the source code restored to plaintext
</pre> 

<!-- grade_id: cell-0e04f341a44584a7, points: 0, grade: false, solution: true -->
```python
def return_encrypted_scytale_code():
    # YOUR CODE HERE
    plaintext =  """scytale = []
    encoded = ""
    i = 0
    k = 0
    while i < len(plaintext):
      loop = []
      while k < N and i < len(plaintext):
        loop.append(plaintext[i])
        i = i + 1
        k = k + 1
      scytale.append(loop)
      k = 0
    for k in range(N):
      for i in range(len(scytale)):
        if k < len(scytale[i]):
          encoded = encoded + scytale[i][k]
    return encoded"""
    return scytale_encode(plaintext,3)
    
def return_decoded_scytale_code(encrypted_code=return_encrypted_scytale_code()):
    # YOUR CODE HERE
    return scytale_cipher_decode(encrypted_code,3)

print(return_decoded_scytale_code())
```

Output:
```
scytale = []
    encoded = ""
    i = 0
    k = 0
    while i < len(plaintext):
      loop = []
      while k < N and i < len(plaintext):
        loop.append(plaintext[i])
        i = i + 1
        k = k + 1
      scytale.append(loop)
      k = 0
    for k in range(N):
      for i in range(len(scytale)):
        if k < len(scytale[i]):
          encoded = encoded + scytale[i][k]
    return encoded
```

<!-- grade_id: cell-bef06a38f2c90ce7, points: 2, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q1.5B --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'ceasarCipherDecode' function..\n ")
assert type(return_encrypted_scytale_code())==str,'output type is incorrect, should be string'
assert type(return_decoded_scytale_code())==str,'output type is incorrect, should be string'
assert type(return_decoded_scytale_code("some_arbitrary_string"))==str,'output type is incorrect, should be string'
print ("good job!\nTests passed. There may be hidden tests too")
```

Output:
```
Testing the implementation of the 'ceasarCipherDecode' function..
 
good job!
Tests passed. There may be hidden tests too
```

# Question 2: Hash tables

The code cell below contains emergency phone numbers of the university (taken from https://in.bgu.ac.il/logistics/Pages/dialing_rules.aspx). A single entity may have several phone numbers.

You will implement a ‘truecaller’ facility: given a phone number, output the name of the entity. For example,
```python
caller = truecaller(t, 4100)
print(f"{caller} called!")
```
```truecaller``` should return “Police”. 

if the number is not found, the function should return None

Q1. Implement 'truecaller' using a Python dictionary


```python
  phonebook = [
    ["Security",
     '61888', '61555', '61444'],
    ["Police",
    '100', '4100'],
    ["Magen David Adom",
    '101', '4101'],
    ["Fire brigade",
    '102', '4102'],
    ["Soroka hospital",
    '4007']
]
    
```

Q2.1A : Implement an auxiliary `make_trucaller_python_dictionary` function that converts the provided `phonebook` list into a Python dictionary
<pre> 
------------
Input parameters:
- phn - a Python list, as `phonebook`, where every item is a list of strings. The first string is the caller's name, the following strings are phone numbers
------------
return value:
Python dictionary
------------
</pre>

<!-- grade_id: cell-87a450fdc22b0d47, points: 0, grade: false, solution: true -->
```python
def make_truecaller_python_dictionary(phn=phonebook):
    # YOUR CODE HERE
    diction = {}
    for i in range(len(phn)):
      for k in range(1,len(phn[i])):
        diction[phn[i][k]] = phn[i][0]
    return diction
```

Q2.1B : Implement a `truecaller1` function, relying on `make_truecaller_python_dictionary`, which takes as input a phone number, and outputs a string with the caller's name as it appears in the phonebook (or returns None if the number couldn't be found).
<pre> 
------------
Input parameters:
- phone_number - a string, containing a phone number
------------
return value:
a string , with the caller's name
------------
</pre>

<!-- grade_id: cell-9e86c1484dcf36f4, points: 0, grade: false, solution: true -->
```python
python_dict= make_truecaller_python_dictionary(phonebook)
def truecaller1(phone_number,caller_dict=python_dict):
    # YOUR CODE HERE
    return caller_dict.get(phone_number)
```

<!-- grade_id: cell-1be6975be3ec4742, points: 12, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q2.1 --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'make_truecaller_python_dictionary' function..\n ")
assert type(make_truecaller_python_dictionary())==dict,'output type is incorrect, should be python dictionary'
print ("Testing the implementation of the 'truecaller1' function..\n ")
assert type(truecaller1('101'))==str,'output type is incorrect, should be string'
print ("good job!\nTests passed. There may be hidden tests too")

```

Output:
```
Testing the implementation of the 'make_truecaller_python_dictionary' function..
 
Testing the implementation of the 'truecaller1' function..
 
good job!
Tests passed. There may be hidden tests too
```


Q2. Implement "truecaller2" using your own implementation of dictionary:

<!-- grade_id: cell-f06f7223c2e94a50, points: 0, grade: false, solution: true -->
```python
#You can add auxiliary functions, according to your discretion
# YOUR CODE HERE
def maket(n):
    # t = []
    # for _ in range(n):
    #     t.append([])
    # return t
    return [list() for _ in range(n)]

def tset(table, key, value, h=hash):
    i = h(key) % len(table)
    entry = table[i]
    for i, (k, v) in enumerate(entry):
        if key==k:
            entry[i] = (key, value)
            return
    entry.append((key, value))
        
def tget(table, key, h=hash):
    i = h(key) % len(table)
    entry = table[i]
    for i, (k, v) in enumerate(entry):
        if key==k:
            return v
    return None
```

Q2.2A : Implement an auxiliary `make_truecaller_own_implement_dictionary` function that converts the provided `phonebook` list into a structure that behaves like a hash table
    
<pre> 
------------
Input parameters:
- phn - a Python list, as `phonebook`, where every item is a list of strings. The first string is the caller's name, the following strings are phone numbers
------------
return value:
a Python list
------------
</pre>

<!-- grade_id: cell-0f5cd96a463d46c9, points: 0, grade: false, solution: true -->
```python
def make_truecaller_own_implement_dictionary(phBook=phonebook):
    # YOUR CODE HERE
    truecaller2 = maket(len(phBook))
    for i in range(len(phBook)):
      for k in range(1,len(phBook[i])):
        tset(truecaller2, phBook[i][k], phBook[i][0])
    return truecaller2
```

Q2.2B : Implement a `truecaller2` function, relying on `make_truecaller_own_implement_dictionary`, which takes as input a phone number, and outputs a string with the caller's name as it appears in the phonebook (or returns None if the number couldn't be found).
<pre> 
------------
Input parameters:
- phone_number - a string, containing a phone number
------------
return value:
a string , with the caller's name
------------
</pre>

<!-- grade_id: cell-a2c907715447ebae, points: 0, grade: false, solution: true -->
```python
make_own_dictionary=make_truecaller_own_implement_dictionary()
def truecaller2(phone_number,own_dictionary=make_own_dictionary):
    # YOUR CODE HERE
    return tget(own_dictionary,phone_number)

print(truecaller2("102",own_dictionary=make_own_dictionary))
```

Output:
```
Fire brigade
```

<!-- grade_id: cell-9bd7a35c7e4305a6, points: 12, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q2.2 --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'make_truecaller_python_dictionary' function..\n ")
assert type(make_truecaller_own_implement_dictionary())==list,'output type is incorrect, should be list'
print ("Testing the implementation of the 'truecaller1' function..\n ")
assert type(truecaller2('102'))==str,'output type is incorrect, should be string'
print ("good job!\nTests passed. There may be hidden tests too")

```

Output:
```
Testing the implementation of the 'make_truecaller_python_dictionary' function..
 
Testing the implementation of the 'truecaller1' function..
 
good job!
Tests passed. There may be hidden tests too
```


3. Implement a hash function `hash_string` for phone numbers. Treat phone numbers as _strings_ of characters. 


<pre> 
------------
Input parameters:
- str_to_hash - a string
------------
return value:
an integer
------------
</pre>

<!-- grade_id: cell-907e0df43371898d, points: 0, grade: false, solution: true -->
```python
def hash_string(str_to_hash):
    # YOUR CODE HERE
    h = 0
    for c in str_to_hash:
        h = 31*h + ord(c)
    return h
```

<!-- grade_id: cell-6074fd7cee4c2b6d, points: 12, grade: true, solution: false -->
```python

# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q2.3 --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'hash_string' function..\n ")
assert type(hash_string("some string"))==int,'output should be an integer'
print ("good job!\nTests passed. There may be hidden tests too")


```

Output:
```
Testing the implementation of the 'hash_string' function..
 
good job!
Tests passed. There may be hidden tests too
```


4. Hashing is called ‘perfect’ if there are no ‘key conflicts‘ — each entry in the hash table contains a single key-value pair. (Remember, we saw in class that if the table size is too small, multiple keys are mapped to the same entry).  What is the minimum table size that creates a perfect hash for the phonebook data with the hash from 2 (built-in hash), and from 3 (custom hash on strings)? Implement functions that return the minimum lengths for each case. 

Implement functions `min_size_for_perfecthash_builtin`,`min_size_for_perfecthash_strhash` that return the correct answers for each type of hash - the built-in Python hash and the function `hash_string` you have implemented above:

<pre> 
------------
Input parameters: none
------------
return value: an integer, the minimum size for the appropriate type of hash
------------
</pre>

<!-- grade_id: cell-a5389da6f18d38bd, points: 0, grade: false, solution: true -->
```python
#You can add auxiliary functions, according to your discretion
# YOUR CODE HERE
def truecaller3(size, h = hash_string, phB = phonebook):
    theTable = maket(size)
    for i in range(len(phB)):
      for k in range(1,len(phB[i])):
        tset(theTable,phB[i][k],phB[i][0],h)
    return theTable
```

<!-- grade_id: cell-cba9f3c0fcdfb195, points: 0, grade: false, solution: true -->
```python
def min_size_for_perfecthash_builtin():
    # YOUR CODE HERE
    erorr = True
    size = 1
    while erorr == True:
      advance = False
      test_table = truecaller3(size,hash)
      for i in test_table:
        if len(i) > 1:
          advance = True
          break
      if advance == True:
        size = size + 1
      else:
        erorr = False
    return size

def min_size_for_perfecthash_strhash():
    # YOUR CODE HERE
    erorr = True
    size = 1
    while erorr == True:
      advance = False
      test_table = truecaller3(size,hash_string)
      for i in test_table:
        if len(i) > 1:
          advance = True
          break
      if advance == True:
        size = size + 1
      else:
        erorr = False
    return size
```

<!-- grade_id: cell-cf67edd4a8ae7aa5, points: 14, grade: true, solution: false -->
```python


# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q2.4 --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the min_size- functions..\n ")
assert type(min_size_for_perfecthash_builtin())==int,'output type should be an integer'
assert type(min_size_for_perfecthash_strhash())==int,'output type should be an integer'
print ("good job!\nTests passed. There may be hidden tests too")

```

Output:
```
Testing the implementation of the min_size- functions..
 
good job!
Tests passed. There may be hidden tests too
```