Instructions:
1. Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE"
2. Do not add or remove cells to the document, inside the current cells you can add helper functions if you wish.
3. Download the file and submit only one file named hw4.ipynb (Do not change file name!)
4. Work alone, you can use any resource that was presented in class!
5. Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Runtime$\rightarrow$Run All).
6. Clear all outputs - Edit$\rightarrow$Clear all outputs


---

# Assignment 4: Differentiation and optimization

```python
import sympy as sym # symbolic differentiation
import jax          # algorithmic differentiation
import jax.numpy as np
import matplotlib.pyplot as plt
```

## Question 1: Differentiation

Function 

$$f(a, b) = \frac 2 b cos(a) \exp \left( - \frac {a^2} {b^2}\right)$$

is given.

1. Derive the partial derivatives of $f(a, b)$ by $a$ and $b$.

    $$\frac {\partial f} {\partial a} = $$
    $$\frac {\partial f} {\partial b} = $$

Implement the function f_partial_derviatives:

<pre>
Input parameters:
a - first symbol
b - second symbol
return value - a tuple (dfa, dfb) such that dfa is the partial derivatives of f by a,
               and dfb is the partial derivative of f by b
</pre> 

<!-- grade_id: cell-fcc26ea16c7fcc6e, points: 0, grade: false, solution: true -->
```python
def f_partial_derviatives_symbol(a, b):
    # YOUR CODE HERE
    f = (2/b)*sym.cos(a)*sym.exp(-((a**2)/(b**2)))
    return (sym.diff(f, a), sym.diff(f, b))
```

<!-- grade_id: cell-910c1caba85f9efe, points: 2, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q1.1 --- Test your implementation:
# ---------------------------
print("Testing the implementation of the 'f_partial_derviatives_symbol' function..\n ")
a=sym.Symbol('a')
b=sym.Symbol('b')
f_deriv_a, f_deriv_b = f_partial_derviatives_symbol(a, b)
assert isinstance(f_deriv_a, sym.Basic)
assert isinstance(f_deriv_b, sym.Basic)
print ("good job!\nSanity tests passed. There are additional hidden tests...")

```

2. Derive the same derivatives by a and b, but using using algorithmic differentiation this time (with `jax`).

Implement the function f_partial_derviatives_algo:

<pre>
Input parameters:
return value - a tuple (dfa, dfb) such that dfa is the partial derivatives of f by a,
               and dfb is the partial derivative of f by b
</pre> 

<!-- grade_id: cell-68ca280308184043, points: 0, grade: false, solution: true -->
```python
def f_partial_derviatives_algo():
    # YOUR CODE HERE
    dfunction1a = jax.grad(function1,0)
    dfunction1b = jax.grad(function1,1)
    return (dfunction1a,dfunction1b)

def function1(a,b):
    return (2/b)*np.cos(a)*np.exp(-((a**2)/(b**2)))
```

<!-- grade_id: cell-bf714cb7381ed64d, points: 2, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q1.2 --- Test your implementation:
# ---------------------------
print("Testing the implementation of the 'f_partial_derviatives_algo' function..\n ")
a=1.
b=1.
f_deriv_a, f_deriv_b = f_partial_derviatives_algo()
assert callable(f_deriv_a)
assert callable(f_deriv_b)
assert f_deriv_a(a, b) < 0
assert f_deriv_b(a, b) > 0
print ("good job!\nSanity tests passed. There are additional hidden tests...")

```

3. Plot the following:

a. One subplot for both:

  * $f(a, 10)$, $\frac {\partial f(a, 10)} {\partial a}$ for range $a \in [-20, 20]$, 
  
b. A seperate subplot below the first subplot for:
  
  * $f(10, b)$, $\frac {\partial f(10, b)} {\partial b}$ for range $b \in [1, 100]$.
  

<!-- grade_id: cell-7c4c505ac0fa2e26, points: 2, grade: true, solution: true -->
```python
# show you plots here:
# YOUR CODE HERE
a = np.linspace(-20, 20, 1000)
f_deriv_a, f_deriv_b = f_partial_derviatives_algo()
plt.plot(a, [function1(a,10) for a in a], label="f(a,10)")
plt.plot(a, [f_deriv_a(a,10) for a in a], label="d f/d a")
plt.legend()
plt.show()
print("                                 ")
b = np.linspace(1, 100, 1000)
plt.plot(b, [function1(10,b) for b in b], label="f(10,b)")
plt.plot(b, [f_deriv_b(10,b) for b in b], label="d f/d b")
plt.legend()
plt.show()
```

4. Implement a function for approximate numerical differentiation, given the difference size $h$.

Implement the function f_partial_derviatives_numeric_at_a_b:

<pre>
Input parameters:
a - first parameter of f
b - second parameter of f
h - the change of the variable to approximate the derivative by
return value - a tuple that contains partial derivatives of f by a and b approximated at (a,b)
</pre> 

<!-- grade_id: cell-1dc99d4eaa406e89, points: 0, grade: false, solution: true -->
```python
def f_partial_derviatives_numeric_at_a_b(a, b, h):
    # YOUR CODE HERE
    fa = (function1(a + h, b) - function1(a, b)) / h
    fb = (function1(a, b + h) - function1(a, b)) / h
    return (fa,fb)
```

<!-- grade_id: cell-f9a81b3184e827ab, points: 2, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q1.4 --- Test your implementation:
# ---------------------------
print("Testing the implementation of the 'f_partial_derviatives_numeric' function..\n ")
a=1.
b=1.
f_deriv_a_at_a_b, f_deriv_b_at_a_b = f_partial_derviatives_numeric_at_a_b(a, b, 1e-12)
assert f_deriv_a_at_a_b < 0
assert f_deriv_b_at_a_b > 0
print ("good job!\nSanity tests passed. There are additional hidden tests...")

```

5. Find the best difference size $h$ for differentiating
   * $f(a, b)$ by $a$.
   * $f(a, b)$ by $b$.
   
The best difference size minimizes the absolute error of numerical differentiation relative to the exact differentiation.
   
Implement the function f_deriv_by_a_optimize_h:

<pre>
Input parameters:
a - first parameter of f
b - second parameter of f
return value - the best difference h to do numeric differentiation of f(a,b) by a
</pre> 

Implement the function f_deriv_by_b_optimize_h:

<pre>
Input parameters:
a - first parameter of f
b - second parameter of f
return value - the best difference h to do numeric differentiation of f(a,b) by b
</pre> 

Non-Mandatory: draw plots of the abolute error by h 
(it's a good practice for you to validate that the results make sense)

* Allowed relative tolerance is up to 100 (so if best h=1e-04, any value between 1e-02 and 1e-06 is allowed)

<!-- grade_id: cell-1bd9e03b2999d29f, points: 0, grade: false, solution: true -->
```python
def f_deriv_by_a_optimize_h(a, b):
    # YOUR CODE HERE
    h = 1
    optimized = 0.1
    numeric = f_partial_derviatives_numeric_at_a_b(a,b,h)
    algo = f_partial_derviatives_algo()[0](a, b)
    opt = abs(algo-numeric[0])
    while numeric[0]!= 0:
      h = h * 0.1
      numeric = f_partial_derviatives_numeric_at_a_b(a,b,h)
      dif = abs(algo-numeric[0])
      if dif < opt:
        optimized = h
        opt = dif
    return optimized 
```

<!-- grade_id: cell-aafc10f1c62a469a, points: 0, grade: false, solution: true -->
```python
def f_deriv_by_b_optimize_h(a, b):
    # YOUR CODE HERE
    h = 1
    optimized = 0.1
    numeric = f_partial_derviatives_numeric_at_a_b(a,b,h)
    algo = f_partial_derviatives_algo()[1](a, b)
    opt = abs(algo-numeric[1])
    while numeric[1]!= 0:
      h = h * 0.1
      numeric = f_partial_derviatives_numeric_at_a_b(a,b,h)
      dif = abs(algo-numeric[1])
      if dif < opt:
        optimized = h
        opt = dif
    return optimized 
```

<!-- grade_id: cell-65a8e984bababd74, points: 2, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q1.5 --- Test your implementation:
# ---------------------------
print("Testing the implementation of the 'f_partial_derviatives_numeric' function..\n ")
assert f_deriv_by_a_optimize_h(4.,1.) < 1e-05
print ("good job!\nSanity tests passed. There are additional hidden tests...")

```

## Question 2: Optimization

### Logistic regression

For a trial group of care bares (https://en.wikipedia.org/wiki/Care_Bears), amount of sugar in blood and the event of sleepiness are given as a list of pairs (sugar, sleep) (1 corresponds to sleepiness):


```python
care_bears = \
[(15.5, 1), (6.5, 0), (49.5, 1), (17.5, 1), (43.0, 1), (42.5, 1), (11.5, 0), 
 (23.5, 0), (17.5, 1), (11.0, 0), (13.5, 0), (9.0, 0), (10.5, 1), (31.0, 0),
 (30.0, 1), (30.0, 1), (47.0, 1), (44.5, 1), (15.5, 0), (46.0, 0), (3.0, 1),
 (42.0, 0), (67.0, 0), (2.5, 1)
]
```

We want to predict care bear sleepiness based on the amount of the sugar. The prediction function is

$$sleep = \frac{sugar}{100} \ge threshold.$$

The loss for this _classification_ problem is called log_loss:

\begin{equation}
\begin{aligned}
& L = -\sum_{i=1}^N (sleep_i \log p_i + (1 - sleep_i) \log (1 - p_i))
\end{aligned}
\end{equation}
where the predicted probability is:
\begin{equation}
\begin{aligned}
& p_i = \frac 1 {1 + \exp(threshold - \frac{sugar_i}{100})}
\end{aligned}
\end{equation}

1. Implement the loss as a function of the threshold.

Implement the function care_bare_classification_loss:

<pre>
Input parameters:
threashold - the treshold of the classification
data - the data in the format of list of tuples
return value - the loss as defined above
</pre> 

<!-- grade_id: cell-da55752637b08a28, points: 0, grade: false, solution: true -->
```python
def care_bare_classification_loss(threshold, data=care_bears):
    # YOUR CODE HERE
    sum = 0
    for i in range(len(data)):
      pi = 1/(1+np.exp(threshold-(data[i][0]/100)))
      value = data[i][1]*np.log(pi) + (1 - data[i][1])*np.log(1-pi)
      sum = sum + value
    return - sum
```

<!-- grade_id: cell-1e51ea7117c6868f, points: 2, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q2.1 --- Test your implementation:
# ---------------------------
print("Testing the implementation of the 'loss' function..\n ")
assert care_bare_classification_loss(0.5) > 0
print ("good job!\nSanity tests passed. There are additional hidden tests...")

```

2. Plot the loss and the derivative of the loss

Implement a function care_bare_classification_loss_deriv

<pre>
Input parameters:
threashold - the treshold of the classification
return value - the derviative of the loss
</pre> 

Now plot the loss and the derivative of the loss by the threshold in the range $threshold \in (0.01, 0.99)$

<!-- grade_id: cell-9e65d966d3fe2745, points: 0, grade: false, solution: true -->
```python
def care_bare_classification_loss_deriv(threshold):
    # YOUR CODE HERE
    dloss = jax.grad(care_bare_classification_loss)
    return dloss(threshold, data=care_bears)

threshold = np.linspace(0.01, 0.99, 100)
plt.plot(threshold, [care_bare_classification_loss(threshold, care_bears) for threshold in threshold], label="loss")
plt.plot(threshold, [care_bare_classification_loss_deriv(threshold) for threshold in threshold], label="dloss")
plt.legend()
```

<!-- grade_id: cell-779160dbb3fa5fbf, points: 2, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q2.2 --- Test your implementation:
# ---------------------------
print("Testing the implementation of the 'loss' function..\n ")
assert care_bare_classification_loss_deriv(0.5) > 0
print ("good job!\nSanity tests passed. There are additional hidden tests...")

```

3. Find the best threshold using gradient descent - use jax for the dervative!

Implement a function loss_best_threshold_gd

<pre>
return value - the best threshold according to the loss function 
</pre> 

<!-- grade_id: cell-252b27d3b325b9e9, points: 0, grade: false, solution: true -->
```python
def loss_best_threshold_gd():
    # YOUR CODE HERE
    return gd(care_bare_classification_loss, 0.1,0.1,0.995,100)

def gd(f, x0, step=0.1, decay=0.995, niter=100):
  """approximates minimum of f starting from x0
  """
  df = jax.grad(f)
  x = x0
  for i in range(niter):
    x -= df(x)*step
    step *= decay
  return x

```

<!-- grade_id: cell-7e98f7a26296840a, points: 2, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q2.3 --- Test your implementation:
# ---------------------------
print("Testing the implementation of the 'loss' function..\n ")
best_threshold = loss_best_threshold_gd()
assert best_threshold > 0
assert care_bare_classification_loss(best_threshold) < care_bare_classification_loss(best_threshold+0.5)
print ("good job!\nSanity tests passed. There are additional hidden tests...")

```

4. Find the best threshold using Newton's method.

Implement a function loss_best_threshold_newton

<pre>
return value - the best threshold according to the loss function 
</pre> 

<!-- grade_id: cell-94bc44c21f3a1b2f, points: 0, grade: false, solution: true -->
```python
def loss_best_threshold_newton():
    # YOUR CODE HERE
    return newton(care_bare_classification_loss, 0.1,10)

def newton(f, x0, niter=10):
  df = jax.grad(f)
  ddf = jax.grad(df)
  x = x0
  for i in range(niter):
    x = x - df(x)/ddf(x)
  return x
```

<!-- grade_id: cell-e03d19aad1688287, points: 2, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q2.4 --- Test your implementation:
# ---------------------------
print("Testing the implementation of the 'loss' function..\n ")
best_threshold = loss_best_threshold_newton()
assert best_threshold > 0
assert care_bare_classification_loss(best_threshold) < care_bare_classification_loss(best_threshold+0.5)
print ("good job!\nSanity tests passed. There are additional hidden tests...")

```

5. Using the best threshold found, how many sleep=1 cases were misclassified?

Implement a function misclassified_sleeps

<pre>
Input parameters:
data - the data in the format of list of tuples
use_newton - True if using Newton, False if using GD
return value - how many sleep=1 cases were missclassified
</pre> 


<!-- grade_id: cell-cfdeadebe7272bf1, points: 0, grade: false, solution: true -->
```python
def misclassified_sleeps(use_newton=False, data=care_bears):
    # YOUR CODE HERE
    threshold = 0
    missclassified = 0
    if use_newton == True:
      threshold = loss_best_threshold_newton()
    else:
      threshold = loss_best_threshold_gd()
    for i in range(len(data)):
      if data[i][1] == 1:
        if data[i][0]/100 < threshold:
          missclassified = missclassified +1
    return missclassified
```

<!-- grade_id: cell-a6bf63dd69b400cc, points: 2, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q2.5 --- Test your implementation:
# ---------------------------
print("Testing the implementation of the 'loss' function..\n ")
missclassified_sleep_count_gd = misclassified_sleeps(False)
missclassified_sleep_count_newton = misclassified_sleeps(True)
assert missclassified_sleep_count_gd >= 0
assert missclassified_sleep_count_newton >= 0
print ("good job!\nSanity tests passed. There are additional hidden tests...")

```